Transcriptome-wide identification of 5-methylcytosine by deaminase and reader protein-assisted sequencing

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1:

In this manuscript, Zhou et al describe a deaminase and reader protein-assisted RNA m5C sequencing method. The general strategy is similar to DART-seq for m6A sequencing, but the difference is that in DART-seq, m6A sites are always followed by C which can be deaminated by fused APOBEC1 to provide a high resolution of m6A sites, while in the case of m5C, no such obvious conserved motifs for m5C sites exist, therefore, the detection resolution is much lower. In addition, the authors used two known m5C binding proteins ALYREF and YBX1 to guide the fused deaminases, but it is not clear whether these two binding proteins can bind most m5C sites and compete with other m5C binding proteins.

Thank you for your kind suggestion. RNA affinity chromatography and mass spectrometry analyses using biotin-labelled oligonucleotides with or without m5C were performed in previous reports (doi:10.1038/cr.2017.55 and doi: 10.1038/s41556-019-0361-y), and the results showed that ALYREF and YBX1 had a more prominent binding ability to m5C -modified oligonucleotides. Moreover, these two m5C -binding proteins are also responsible for mRNA m5C binding, so we chose to use their ability to bind targeted m5C to construct a DRAM detection system in anticipation of transcriptome-wide m5C detection. We hope to propose a suitable detection strategy for RNA m5C, and there will certainly be room for optimization of the DRAM system in the future with more in-depth studies of m5C binding proteins. We have discussed the above issue in lines 75-82 and 315-318.

It is well known that two highly modified m5C sites exist in 28S RNA and many m5C sites exist in tRNA, the authors should validate their methods first by detecting these known m5C sites and evaluate the possible false positives in rRNA and tRNA.

Thank you for your kind suggestion. We attempted PCR amplification of sequences flanking m5C sites 3782 and 4447 on 28S rRNA, as well as multiple m5C sites on tRNA, including m5C48 and m5C49 on tRNAVal, m5C48 and m5C49 on tRNAAsp, and m5C48 on tRNALys.

However, Sanger sequencing revealed no valid mutations, which was implemented in Figure S3. We believe this outcome indicates that the DRAM system is more suited for transcriptome-wide m5C detection of mRNAs. This is supported by current reports that ALYREF and YBX1 are responsible for the m5C-binding proteins of mRNAs (doi:10.1038/cr.2017.55 and doi: 10.1038/s41556-019-0361-y). The above results and descriptions were added to lines 136-143.

In mRNA, it is not clear what is the overlap between the technical replicates. In Figures 4A and 4C, they detected more than 10K m5C sites, and most of them did not overlap with sites uncovered by other methods. These numbers are much larger than expected and possibly most of them are false positives.

Thank you for your kind suggestion. We observed significant overlap between the technical repeats by comparing the data across biological repeats, as shown in Figure S4C and described in lines 174-175. We considered m5C modification in a region only when editing events were detected in at least two biological replicates, ensuring a high-stringency screening process (details seen in the revised method in lines 448-455 and Figure 3F). With more in-depth research into m5C readers, we aim to achieve more accurate detection in the future.

Besides, it is not clear what is the detection sensitivity and accuracy since the method is neither single base resolution nor quantitative.

Thank you for your suggestion. As shown in Figure 3G, we found that the editing window of the DRAM system exhibited enrichment of approximately 20 bp upstream and downstream of the m5C site. Previous reports identified Type I m5C sites, which tend to have a downstream "NGGG" motif, and Type II m5C sites, which often contain a downstream "UCCA" motif. However, these m5C motifs do not fully characterize all m5C sites, and their presence downstream of an m5C site is not guaranteed (doi: 10.1038/s41594-019-0218-x). This limitation complicates single-base resolution analysis by the DRAM system. Nevertheless, we believe that with further exploration of m5C sequence features, precise single-base resolution detection can be achieved in the future. This point is also discussed in lines 314-322.

Regarding the quantitative level of the assay, we conducted additional experiments by progressively reducing the expression levels of the fusion proteins. Sanger sequencing revealed that the editing efficiency of A-to-G and C-to-U within the m5C region significantly decreased as fusion protein expression diminished (Figure S9). These findings suggest that the DRAM system's transfection efficiency is concentration-dependent and that the ratio of editing efficiency to transfection efficiency could aid in the quantitative analysis of m5C using the DRAM system. The relative results were supplemented in Figure S9 and discussed in lines 263-271.

There are no experiments to show that the detected m5C sites are responsive to the writer proteins such as NSUN2 and NSUN6, and the determination of the motifs of these writer proteins.

Thank you for your kind suggestion. We have performed a motif enrichment analysis based on the sequences spanning 10 nt upstream and downstream of DRAM-editing sites. The relative results of this analysis were supplemented in Figure S4D and lines 168-171. Unfortunately, we did not identify any clear sequence preferences for the m5C sites catalyzed by the methyltransferases NSUN2 and NSUN6, which have previously been associated with “G”-rich sequences and the “CUCCA” motif. This limitation is mainly due to the DRAM detection system’s inability to achieve single-base resolution for m5C detection, which is also explained in the above response.

Reviewer #2:

(1) The use of two m5C reader proteins is likely a reason for the high number of edits introduced by the DRAM-Seq method. Both ALYREF and YBX1 are ubiquitous proteins with multiple roles in RNA metabolism including splicing and mRNA export. It is reasonable to assume that both ALYREF and YBX1 bind to many mRNAs that do not contain m5C.

To substantiate the author's claim that ALYREF or YBX1 binds m5C-modified RNAs to an extent that would allow distinguishing its binding to non-modified RNAs from binding to m5C-modified RNAs, it would be recommended to provide data on the affinity of these, supposedly proven, m5C readers to non-modified versus m5C-modified RNAs. To do so, this reviewer suggests performing experiments as described in Slama et al., 2020 (doi: 10.1016/j.ymeth.2018.10.020). However, using dot blots like in so many published studies to show modification of a specific antibody or protein binding, is insufficient as an argument because no antibody, nor protein, encounters nanograms to micrograms of a specific RNA identity in a cell. This issue remains a major caveat in all studies using so-called RNA modification reader proteins as bait for detecting RNA modifications in epitranscriptomics research. It becomes a pertinent problem if used as a platform for base editing similar to the work presented in this manuscript.

We thank the reviewer for the valuable suggestion. Previous studies have shown that while ALYREF and YBX1 can bind mRNAs without the m5C modification, their binding affinity for m5C-modified oligonucleotides is significantly higher than for unmethylated controls. This has been demonstrated through experiments such as in vitro tractography, electrophoretic mobility shift assay (EMSA) (doi:10.1038/cr.2017.55), and UHPLC-MRM-MS/MS. Additionally, isothermal titration calorimetry measurements and PAR-CLIP experiments have shown that mutations in the key amino acids responsible for m5C binding in ALYREF and YBX1 result in a significant reduction in their ability to m5C (doi: 10.1038/s41556-019-0361-y).

Although Me-RIP analysis was unsuccessful in our laboratory, likely due to the poor specificity of the m5C antibody, we alternatively performed RNA pulldown experiments. These experiments verified that the ability of DRAMmut-expressing proteins to bind RNA with m5C modification was virtually absent compared to DRAM-expressing proteins, while their binding ability with non-modified RNA was not significantly affected. The relative RNA pulldown results were supplemented in Figure S1E, S1F and lines 110-111. Therefore, we believe that by integrating DRAMmut group, our DRAM system could effectively exclude the false-positive mutations caused by unspecific binding of DRAM’s reader protein to non-m5C-modified mRNAs.

(2) Since the authors use a system that results in transient overexpression of base editor fusion proteins, they might introduce advantageous binding of these proteins to RNAs. It is unclear, which promotor is driving construct expression but it stands to reason that part of the data is based on artifacts caused by overexpression. Could the authors attempt testing whether manipulating expression levels of these fusion proteins results in different editing levels at the same RNA substrate?

Thank you for pointing this out. To investigate how different expression levels of these proteins influence A-to-G and C-to-U editing within the same m5C region, we conducted a gradient transfection using plasmid concentrations of 1500 ng, 750 ng and 300 ng. This approach allowed us to progressively reduce the expression levels of the fusion proteins. Sanger sequencing revealed that the editing efficiency of A-to-G and C-to-U within the m5C region significantly decreased as fusion protein expression diminished. These findings suggest that the transfection efficiency of the DRAM system is concentration-dependent and that the ratio of editing efficiency to transfection efficiency may assist in the quantitative analysis of m5C using the DRAM system. The relative results and hypotheses were added and discussed in Figure S9 and lines 263-271 of the revised manuscript.

(3) Using sodium arsenite treatment of cells as a means to change the m5C status of transcripts through the downregulation of the two major m5C writer proteins NSUN2 and NSUN6 is problematic and the conclusions from these experiments are not warranted. Sodium arsenite is a chemical that poisons every protein containing thiol groups. Not only do NSUN proteins contain cysteines but also the base editor fusion proteins. Arsenite will inactivate these proteins, hence the editing frequency will drop, as observed in the experiments shown in Figure 5, which the authors explain with fewer m5C sites to be detected by the fusion proteins.

Thank you for pointing this out. We used bisulfite sequencing PCR to determine that the m5C levels in RPSA and AP5Z1 were significantly reduced after sodium arsenite treatment. This was followed by a significant decrease in editing frequency detected by the DRAM system in sodium arsenite-treated samples compared to untreated samples. This reduction aligns with the decreased editing efficiency observed in methyltransferase-deficient cells (as shown in Figures 2G and 2H), which initially convinced us that these results reflected the DRAM system's ability to monitor dynamic changes in m5C levels.

However, as the reviewer pointed out, sodium arsenite treatment could potentially inactivate the fusion proteins, leading to the observed reduction in editing efficiency. This possibility has not been conclusively ruled out in our current experiments. Optimizing this validation may require the future development of more specific m5C inhibitors. In light of this, we have revised our previous results and conclusions in lines 235-244, and discussed these points in lines 308-315.

(4) The authors should move high-confidence editing site data contained in Supplementary Tables 2 and 3 into one of the main Figures to substantiate what is discussed in Figure 4A. However, the data needs to be visualized in another way than an Excel format. Furthermore, Supplementary Table 2 does not contain a description of the columns, while Supplementary Table 3 contains a single row with letters and numbers.

Thank you for your kind suggestion. We have visualized the data from Supplementary Tables 2 and 3 into Figure 3F, presenting it as a screening flowchart for high-confidence editing sites. In Supplementary Table 3, we have displayed only the DRAM-mutated genes, which is why it contains a single row with letters and numbers. As requested, we have included descriptions of each column and reorganized the Supplementary table 2 and 3 accordingly.

(5) The authors state that "plotting the distribution of DRAM-seq editing sites in mRNA segments (5'UTR, CDS, and 3'UTR) highlighted a significant enrichment near the initiation codon (Figure 3F).", which is not true when this reviewer looks at Figure 3F.

Thank you for your kind suggestion, and we replaced the expression of " near the initiation codon" with "in the CDS" in lines 192-193.

(6) The authors state that "In contrast, cells expressing the deaminase exhibited a distinct distribution pattern of editing sites, characterized by a prevalence throughout the 5'UTR.", which is not true when this reviewer looks at Figure 3F.

Thank you for your kind suggestion. This distribution was actually characterized by a prevalence throughout the "3'UTR", but not "5'UTR". We have also made the necessary changes in lines 193-195.

(7) The authors claim in the final conclusion: "In summary, we developed a novel deaminase and reader protein assisted RNA m5C methylation approach...", which is not what the method entails. The authors deaminate As or Us close to 5mC sites based on the binding of a deaminase-containing protein.

Thank you for your kind suggestion, and we have made the necessary changes in lines 331-334.

(8) The authors claim that "The data supporting the findings of this study are available within the article and its Supplementary Information." However, no single accession number for the deposited sequencing data can be found in the text or the supplementary data. Without the primary data, none of the claims can be verified.

Thank you for pointing this out. The sequencing data from this study has already been deposited to the GEO database (GEO assession number: GSE254194, GEO token：ororioukbdqtpcn), and we will ensure it is made publicly available in a timely manner.

(a) To underscore point (1), a recent publication (https://doi.org/10.1038/s41419-023-05661-y) reported: "To further identify the potential mRNAs regulated by ALYREF, we performed RNA-seq analysis in control or ALYREF knockdown T24 cells. After knockdown of ALYREF, 143 mRNAs differentially expressed, including 94 downregulated mRNAs (NC reads >100, |Fold change | >1.5, P-value <0.05). Functional enrichment analysis using Kyoto Encyclopedia of Genes and Genomes (KEGG) indicated that regulated mRNAs by ALYREF are chiefly enriched in canonical cancer-related pathways (Fig. S4A), including TGF-β signaling, MAPK signaling, and NF-κB signaling, strongly supporting the oncogenic function of ALYREF in tumor progression. Among these 94 downregulated genes, 11 mRNA showed a significant reduction in m5C methylation after NUSN2 silencing in T24 cells, combined with previously transcriptome-wide RNA-BisSeq data of T24 cells [21] (Fig. 4A)."

These results translate into 94 mRNAs are regulated by ALYREF in bladder cancer-derived cells. From those, very few (11) mRNA identities respond to NSUN2-dependent RNA methylation mediated by ALYREF binding.The question then arises, is that number sufficient to claim that ALYREF is a m5C-binding protein?

And if so, how does the identification of 10.000+ edits by DRAM-Seq compare with the 94 mRNAs that are regulated by ALYREF? Were these 94 mRNAs identified by DRAM-Seq.

Thank you for your kind suggestion. Previous reports by Yang et al. ( doi: 10.1038/cr.2017.55), including the literature you refer to, have detailed the close relationship between ALYREF and m5C modification, and the ALY/REF export factor (ALYREF) was identified as the first nuclear m5C reader, and it was demonstrated that many mRNAs are regulated by ALYREF, and is therefore considered to be an m5C-binding protein.

As required, by comparing the DRAM-edited mRNAs with the reported 94 mRNAs, we found that only 55.32% of the 94 mRNAs regulated by ALYREF could be detected by the DRAM system. This indicates that the DRAM system specifically targets certain mRNAs, as illustrated in Figure S4E. The relevant results were described and discussed in lines 175-179.

(b) Line 123:

"The deep sequencing results showed that the deamination rates of RPSA and SZRD1 were 75.5% and 27.25%, respectively. (Fig. 2A, B)."

The Figure shows exactly the opposite of bisulfite-mediated deamination. These are the cytosines that were not deaminated by the chemical treatment and therefore can be sequenced as cytosines and not thymidines. Hence, the term deamination rate is wrong.

Thank you for your kind suggestion. We have made the necessary change in lines 129-130 to change the deamination rates to m⁵C fraction.

(c) Line 157:

"DRAM-seq analysis further confirmed that DRAM was detected in an m5C-dependent manner, with minimal mutations in AP5Z1 and RPSA mRNAs in methyltransferase knockout cells compared to wild-type cells (Fig. 3C, D)."

There is no indication of what the authors mean by minimal mutation in these Figures. The term "minimal mutation" should be reconsidered as well.

Thank you for your kind suggestion. We intended to express that "Mutations in AP5Z1 and RPSA mRNA are reduced in methyltransferase-deficient cells." There was an issue with the initial formulation, and we have made the necessary changes in lines 165-167.

(d) Line 167:

"To further delineate the characteristics of the DRAM-seq data, we compared the distribution of DRAM-seq editing sites within the gene structure, specifically examining their occurrences in the 5'untranslated region (5'UTR), 3' untranslated region (3'UTR), CDS and ncRNA."

Which part of a coding RNA is meant by "ncRNA"?

Thank you for pointing this out. This was actually the Intergenic or Intron region, but not ncRNA. We have also corrected this labelling in Figure 3G and lines 186-189 of the revised manuscript.

(e) Line 189:

"Subsequently, we assessed the capacity of DRAM-seq to detect m5C on a transcriptome-wide scale, comparing its performance to BS-seq that have been previously reported with great authority."

The term "great authority" is not a scientific term. Please, remove adulation to senior authors.

Thank you for your kind suggestion. We removed this unsuitable expression and made the necessary changes in lines 207-208.

(f) Line 233:

"Several experiments have highlighted the requirement of 100-500 ng of RNA for m5C-RIP-seq, while BS-seq necessitates an even more demanding 500-750 μg of RNA21,25,61."

This reviewer doubts that RNA bisulfite sequencing required half to one mg of RNA input. Please, check these references.

Thank you for your kind suggestion. According to the references, we corrected μg to ng and made the necessary changes in lines 251-252.

(g) Line 247:

"Several experiments have highlighted the requirement of 100-500 ng of RNA for m5C-RIP-seq, while BS-seq necessitates an even more demanding 500-750 μg of RNA21,25,61."

This reviewer doubts that RNA bisulfite sequencing requires half to one mg of RNA input. Please, check these references.

Thank you for your kind suggestion. According to the references, we corrected μg to ng and made the necessary changes in lines 251-252.

(h) Line 292:

"Since m5C lacks a fixed motif, DRAM has an apparent limitation in achieving single-base resolution for detecting m5C."

m5C deposition by NSUN2 and NSUN6 occurs in particular motifs that were coined Type I and II motifs. Hence, this statement is not correct.

Thank you for your kind suggestion. Previous reports identified Type I m5C sites, which tend to have a downstream "NGGG" motif, and Type II m5C sites, which often contain a downstream "UCCA" motif. However, these m5C motifs do not fully characterize all m5C sites, and their presence downstream of an m5C site is not guaranteed (doi: 10.1038/s41594-019-0218-x ). Therefore, we have corrected the expression “fixed motif” to “fixed base composition for characterizing all m5C modification sites” in lines 317.

(i) Line 390:

"1 μl of total cellular RNA was used for sequencing library gene..."

1 uL does not allow us to deduce which RNA mass was used for cDNA synthesis.

Thank you for your kind suggestion. According to our cDNA synthesis protocol, we corrected “1μl” to “1μg” in lines 422-423.

(j) Line 405:

"...was assessed on the Agilent 5400 system (Agilent, USA) and quantified by QPCR (1.5 nM)"

What does the 1.5 nM refer to in this sentence?

Thank you for your kind suggestion. Here, "1.5nM" means that the concentration of the constructed library should be no less than 1.5nM. We have also revised this expression in the methods in lines 436-438.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: 5-Methylcytosine (m5C) is one of the major post-transcriptional modifications in mRNA and is highly involved in the pathogenesis of various diseases. However, the capacity of existing assays for accurately and comprehensively transcriptome-wide m5C mapping still needs improvement. Here, we develop a detection method named DRAM (deaminase and reader protein assisted RNA methylation analysis), in which deaminases (APOBEC1 and TadA-8e) are fused with m5C reader proteins (ALYREF and YBX1) to identify the m5C sites through deamination events neighboring the methylation sites. This antibody-free and bisulfite-free approach provides transcriptome-wide editing regions which are highly overlapped with the publicly available BS-seq datasets and allows for a more stable and comprehensive identification of the m5C loci. In addition, DRAM system even supports ultra-low input RNA (10ng) and monitor the dynamic accumulation of cellular m5C. We anticipate that the DRAM system could pave the way for uncovering further biological functions of m5C modifications.

Revised version: 5-Methylcytosine (m5C) is one of the major post-transcriptional modifications in mRNA and is highly involved in the pathogenesis of various diseases. However, the capacity of existing assays for accurately and comprehensively transcriptome-wide m5C mapping still needs improvement. Here, we develop a detection method named DRAM (deaminase and reader protein assisted RNA methylation analysis), in which deaminases (APOBEC1 and TadA-8e) are fused with m5C reader proteins (ALYREF and YBX1) to identify the m5C sites through deamination events neighboring the methylation sites. This antibody-free and bisulfite-free approach provides transcriptome-wide editing regions which are highly overlapped with the publicly available BS-seq datasets and allows for a more stable and comprehensive identification of the m5C loci. In addition, DRAM system even supports ultra-low input RNA (10ng) and monitors the dynamic accumulation of cellular m5C. We anticipate that the DRAM system could pave the way for uncovering further biological functions of m5C modifications.

Introduction

Earlier version: Epigenetics refers to stable inheritance without changing the basic sequence of DNA, involving various forms such as DNA methylation, histone modification and RNA modification. In recent years, RNA sequencing technology has boosted research on RNA epigenetics. More than 170 RNA modifications have been identified, mainly including m6A, m5C, m1A, m7G and others 1,2. Notably, RNA m5C methylation represents a crucial post-transcriptional modification observed across different RNA types, such as tRNA, mRNA, rRNA, vault RNA, microRNA, long non-coding RNA and enhancer RNA3-8. Numerous studies have revealed multiple molecular functions of m5C in numerous key stages of RNA metabolism, such as mRNA stability, translation, nuclear export5,9-13. And the dynamic alterations of m5C play integral roles in many physiological and pathological processes, such as early embryonic development14, neurodevelopmental disorder15,16 and multifarious tumorigenesis and migration17-20. Moreover, this modification significantly contributes to the regulation of gene expression5,9-13,17. Therefore, the detection of m5C sites appears to be essential for understanding its underlying effects on cellular function and disease states.

Revised version: Epigenetics refers to stable inheritance without changing the basic sequence of DNA, involving various forms such as DNA methylation, histone modification and RNA modification. In recent years, RNA sequencing technology has boosted research on RNA epigenetics. More than 170 RNA modifications have been identified, mainly including m6A, m5C, m1A, m7G and others 1,2. Notably, RNA m5C methylation represents a crucial post-transcriptional modification observed across different RNA types, such as tRNA, mRNA, rRNA, vault RNA, microRNA, long non-coding RNA and enhancer RNA3–8. Numerous studies have revealed multiple molecular functions of m5C in numerous key stages of RNA metabolism, such as mRNA stability, translation, and nuclear export5,9–13. The dynamic alterations of m5C play integral roles in many physiological and pathological processes, such as early embryonic development14, neurodevelopmental disorder15,16 and multifarious tumorigenesis and migration17–20. Moreover, this modification significantly contributes to the regulation of gene expression5,9–13,17. Therefore, the detection of m5C sites appears to be essential for understanding their underlying effects on cellular function and disease states.

Earlier version: With the recent advances in sequencing techniques, several high-throughput assays have been developed for qualitative or quantitative analysis of m5C. To date, bisulfite-sequencing (BS-seq) has been proved to be the gold standard method for RNA m5C methylation analysis5,21,22. This approach chemically deaminates unmethylated cytosine to uracil, while kept methylated cytosine unchanged. The m5C methylation sites can be identified by subsequent library construction and sequencing. However, bisulfite treatment of BS-seq is extremely detrimental to RNA, thus resulting in unstable detection of m5C in low abundance RNA or highly structured RNA, which directly affects the confidence of results23,24. Another major type of global m5C analysis depends on antibody-assisted immunoprecipitation of m5C methylated RNAs, such as m5C-RIP-seq25-27, AZA-IP-seq28 or miCLIP-seq7. These methods are unable to recognize methylation on mRNAs with low abundance and secondary structure. Moreover, these methods are highly dependent on antibody specificity, which usually lead to unspecific binding of RNA and low amount of m5C-modified regions. Moreover, TAWO-seq, originally developed for the identification of hm5C, is also capable for m5C analysis, but it highly depends on the oxidation efficiency of perovskite, which usually cause false positives and unstable conversion 29,30. Furthermore, the emerging third generation sequencing, such as Nanopore-seq, can directly map m5C by tracking the characteristic changes of bases, but it still faces challenges of a high error rate31-33.These together largely hampers its wide application on transcriptome profiling of m5C (Supplementary Table 1). Hence, there is an urgent need for a simple, efficient, sensitive, and antibody-independent method for global m5C detection.

Revised version: With the recent advances in sequencing techniques, several high-throughput assays have been developed for qualitative or quantitative analysis of m5C. To date, bisulfite-sequencing (BS-seq) has been proven to be the gold standard method for RNA m5C methylation analysis5,21,22. This approach chemically deaminates unmethylated cytosine to uracil, while keeping methylated cytosine unchanged. The m5C methylation sites can be identified by subsequent library construction and sequencing. However, bisulfite treatment of BS-seq is extremely detrimental to RNA, thus resulting in unstable detection of m5C in low abundance RNA or highly structured RNA, which directly affects the confidence of results23,24. Another major type of global m5C analysis depends on antibody-assisted immunoprecipitation of m5C methylated RNAs, such as m5C-RIP-seq25–27, AZA-IP-seq28 or miCLIP-seq7. These methods are unable to recognize methylation on mRNAs with low abundance and secondary structure. Moreover, these methods are highly dependent on antibody specificity, which usually leads to unspecific binding of RNA and a low amount of m5C-modified regions. Moreover, TAWO-seq, originally developed for the identification of hm5C, is also capable of m5C analysis, but it highly depends on the oxidation efficiency of perovskite, which usually causes false positives and unstable conversion 29,30. Furthermore, the emerging third-generation sequencing, such as Nanopore-seq, can directly map m5C by tracking the characteristic changes of bases, but it still faces challenges of a high error rate31–33. These together largely hamper its wide application on transcriptome profiling of m5C (Supplementary Table 1). Hence, there is an urgent need for a simple, efficient, sensitive, and antibody-independent method for global m5C detection.

Removed in the revised version: The RNA-binding protein ALYREF is the initially recognized nuclear m5C reader that binds directly to m5C sites in mRNA and play key roles in promoting mRNA nuclear export or tumor progression5. Another well-known m5C reader, YBX1, can also recognize m5C-modified mRNA through its cold-shock domain, and participates in a variety of RNA-dependent events such as mRNA packaging, mRNA stabilization and translational regulation9,18. YBX1 can preferentially recognize mRNAs with m5C modifications via key amino acids W65-N70 (WFNVRN)18, while K171 is essential for the specific binding of ALYREF to m5C sites 5.

Earlier version: Nucleic acid deaminases, primarily categorized as cytosine deaminases and adenine deaminases, are zinc-dependent enzymes which facilitate the deamination of cytosine or adenine within DNA or RNA substrates34. APOBEC1, an evolutionarily conserved family member of APOBEC proteins, can specifically catalyze the deamination of cytosine in single stranded RNA (ssRNA) or DNA (ssDNA) to uracil35-37. TadA8e is an adenine deaminase optimized through re-engineering of TadA and it induces conversion of adenine to inosine (eventually read as guanine by transcriptases) in ssRNA or ssDNA38,39. APOBEC1 and TadA8e, with their prominent deamination efficiency, have been employed for the development of precise and efficient base editors such as CBE and ABE8e, which find widespread application in studies related to genome editing37,38.

Revised version: The RNA-binding protein ALYREF is the initially recognized nuclear m5C reader that binds directly to m5C sites in mRNA and plays key roles in promoting mRNA nuclear export or tumor progression5. Another well-known m5C reader, YBX1, can also recognize m5C-modified mRNA through its cold-shock domain and participates in a variety of RNA-dependent events such as mRNA packaging, mRNA stabilization and translational regulation9,18. RNA affinity chromatography and mass spectrometry analyses using biotin-labelled oligonucleotides with or without m5C were performed in previous reports, which indicated that ALYREF and YBX1 had a more prominent binding ability to m5C-modified oligonucleotides5,18. YBX1 can preferentially recognize mRNAs with m5C modifications via key amino acids W65-N70 (WFNVRN)18, while K171 is essential for the specific binding of ALYREF to m5C sites 5. Previous studies have shown that mutations in key amino acids responsible for recognising m5C binding in ALYREF and YBX1 lead to a significant reduction in their binding levels to m5C-containing oligonucleotides5,18. Nucleic acid deaminases, primarily categorized as cytosine deaminases and adenine deaminases, are zinc-dependent enzymes which facilitate the deamination of cytosine or adenine within DNA or RNA substrates34. APOBEC1, an evolutionarily conserved family member of APOBEC proteins, can specifically catalyze the deamination of cytosine in single-stranded RNA (ssRNA) or DNA (ssDNA) to uracil35–37. TadA8e is an adenine deaminase optimized through re-engineering of TadA and it induces conversion of adenine to inosine (eventually read as guanine by transcriptases) in ssRNA or ssDNA38,39. APOBEC1 and TadA8e, with their prominent deamination efficiency, have been employed for the development of precise and efficient base editors such as CBE and ABE8e, which find widespread application in studies related to genome editing37,38.

Here we aim to establish a deaminase and m5C reader-assisted RNA methylation sequencing approach (DRAM-seq), which identifies the m5C sites through reader-mediated recognitions and deaminase-mediated point mutations neighboring the m5C methylation sites. This bisulfite-free and antibody-free method is anticipated to provide more comprehensive and cost-effective transcriptome-wide detection of m5C methylation, which may better assist on exploring its further regulatory mechanisms.

Results

Development of DRAM system for m5C detection

Earlier version: Our sequencing platform is inspired by the concept of the m6A DART-seq assay, in which C near the m6A site is converted into U without affecting sequences near non-m6A sites40. Therefore, we hypothesized that, by utilizing the targeted binding of m5C readers, deaminase can be recruited to achieve deamination of cytosine or adenine in the vicinity of the m5C sites on single-stranded RNA, thereby facilitating the detection of the m5C site. This approach was named DRAM (deaminase and m5C reader-assisted RNA methylation sequencing). As RNA-binding proteins, ALYREF and YBX1 also could bind to RNAs without m5C modification5,18. In order to exclude the false-positive detection of DRAM due to the non-m5C specific binding of ALYREF and YBX1, knockout of W65-N70 (WFNVRN) amino acids in YBX1and K171A mutation in ALYREF were introduced seperately, resulting in the DRAMmut system (Fig. S1, Supporting Information). To confirm the recognition of m5C site by DRAM system, DRAM, DRAMmut and Deaminase system were transfected into the human HEK293T cells, respectively. Finally the m5C regions were identified by detecting the C-to-U/A-to-G transition, which only occur in the presence of the DRAM system, but not in the presence of the corresponding DRAMmut and Deaminase system (Fig. 1A).

Revised version: Our sequencing platform is inspired by the concept of the m6A DART-seq assay, in which C near the m6A site is converted into U without affecting sequences near non-m6A sites40. Therefore, we hypothesized that, by utilizing the targeted binding of m5C readers, deaminase can be recruited to achieve deamination of cytosine or adenine in the vicinity of the m5C sites on single-stranded RNA, thereby facilitating the detection of the m5C site. This approach was named DRAM (deaminase and m5C reader-assisted RNA methylation sequencing). As RNA-binding proteins, ALYREF and YBX1 also could bind to RNAs without m5C modification5,18. To exclude the false-positive detection of DRAM due to the non-m5C specific binding of ALYREF and YBX1, knockout of W65-N70 (WFNVRN) amino acids in YBX1and K171A mutation in ALYREF were introduced seperately, resulting in the DRAMmut system (Fig. S1A-S1D). DRAMmut exhibits substantially reduced binding to m5C in RNA pulldown assays (Fig. S1D-S1F). To confirm the recognition of m5C site by DRAM system, DRAM, DRAMmut and Deaminase system were transfected into the human HEK293T cells, respectively. Finally, we considered the presence of m5C modification in the vicinity only if the deamination changes produced under DRAM induction were significantly different from those produced under DRAMmut or Deaminase induction (Fig. 1A).

Development of DRAM system for m5C detection.

(A) Schematic diagram of the DRAM assay. DRAM, DRAMmut and Deaminase system were transfected into HEK293T cells separately. After DRAM transfection, the deaminase was directed by m5C reader to the vicinity of the m5C site and induce C-to-U/A-to-G mutations, whereas transfection of the DRAMmut or Deaminase system failed to effectively induce similar mutations due to the absence of the m5C-recognition-binding domain.

(B) The overall design of DRAM, DRAMmut and Deaminase system.

Previous studies have indicated that there is no uniform intrinsic signature motif sequence that can characterize all m5C sites5,26,41,42. To comprehensively detect the m5C loci, the readers of m5C (ALYREF and YBX1) were separately fused to the C-terminus of the deaminases (APOBEC1 and TadA-8e), namely DRAM-ABE and DRAM-CBE system (Fig.1B).

DRAM detection system is assayed in an m5C-dependent form

Earlier version: To confirm the recognition of m5C site by DRAM system, DRAM, DRAMmut and Deaminase were transfected into the human HEK293T cells, respectively. Two previously reported m5C sites in RPSA and AP5Z1 were selected for the analysis 5,21, and their methylation status were verified by bisulfite sequencing PCR. The deep sequencing results showed that the deamination rates of RPSA and SZRD1 were 75.5% and 27.25%, respectively. (Fig.2A and B). Sanger sequencing following RT-PCR was then performed to determine the editing of neighbouring m5C sites by DRAM system in these two mRNA. Notably, adenine close to the m5C site in RPSA mRNA was mutated into guanine, resulting in an A-to-G editing rate of 14.7% by DRAM-ABE, whereas this was rarely observed with TadA-8e or DRAMmut-ABE (Fig.2C). DRAM-CBE induced C to U editing in the vicinity of the m5C site in AP5Z1 mRNA, with 13.6% C-to-U editing, while this effect was significantly reduced with APOBEC1 or DRAMmut-CBE (Fig.2D). Taken together, the fusion of m5C reader and deaminase can effectively and selectively deaminate cytosine/adenine in the vicinity of the m5C sites.

Revised version: To confirm the recognition of m5C site by DRAM system, DRAM, DRAMmut and Deaminase were transfected into the human HEK293T cells, respectively. To evaluate candidate DRAM constructs within a cellular environment, we performed fluorescence microscopy to analyze the expression of DRAM. The results showed that DRAM-ABE and DRAM-CBE were properly expressed in HEK293T cells (Fig. S2A-S2B). In addition, flow cytometry displayed ∼60% of cells were GFP-positive (Fig. S2C). Two previously reported m5C sites in RPSA and AP5Z1 were selected for the analysis 5,21, and their methylation status was verified by bisulfite sequencing PCR. The deep sequencing results showed that the m[C fraction of RPSA and SZRD1 was 75.5% and 27.25%, respectively (Fig.2A and B). Sanger sequencing following RT-PCR was then performed to determine the editing of neighbouring m5C sites by DRAM system in these two mRNA. Notably, adenine close to the m5C site in RPSA mRNA was mutated into guanine, resulting in an A-to-G editing rate of 14.7% by DRAM-ABE, whereas this was rarely observed with TadA-8e or DRAMmut-ABE (Fig.2C). DRAM-CBE induced C to U editing in the vicinity of the m5C site in AP5Z1 mRNA, with 13.6% C-to-U editing, while this effect was significantly reduced with APOBEC1 or DRAMmut-CBE (Fig.2D). Subsequently, in order to investigate whether the DRAM system can detect other types of RNA, such as tRNA, 28S rRNA, or others, we performed PCR amplification of the flanking sequences of the m5C sites 3782 and 4447 on 28S rRNA and several m5C sites on tRNA, such as the m5C48 and m5C49 sites of tRNAVal, the m5C48 and m5C49 sites of tRNAAsp, and the m5C48 site of tRNALys. But Sanger sequencing showed that there was no valid A-to-G/C-to-U mutation detected, which is most likely due to the fact that ALYTEF and YBX1 are mainly responsible for the mRNA m5C binding proteins, and thus the DRAM system is more suitable for the mRNA m5C detection (Fig. S3). Taken together, the fusion of m5C reader and deaminase can effectively and selectively deaminate cytosine/adenine in the vicinity of the mRNA m5C sites.

DRAM detection system was assayed in an m5C -dependent form.

Earlier version: (A, B) Two m5C sites from RPSA (A) and AP5Z1 (B) mRNA detected by deep sequencing of bisulfite sequencing PCR in HEK293T cells. The m5C sites are highlighted by red color. The deamination rates of RPSA and AP5Z1 were 75.5% and 27.25% (The number of reads is greater than 1000).

Revised version: (A, B) Two m5C sites from RPSA (A) and AP5Z1 (B) mRNA detected by deep sequencing of bisulfite sequencing PCR in HEK293T cells. The m5C sites are highlighted by red color. The mlC fraction of RPSA and AP5Z1 were 75.5% and 27.25% (The number of reads is greater than 1000).

(C, D) Sanger sequencing following RT-PCR verified two m5C sites from RPSA (C) and AP5Z1

(D) mRNAs in DRAM-transfected HEK293T cells, respectively. HEK293T cells only expressing DRAMmut or Deaminase were served as negative controls. The left panel illustrates the location of DRAM induced mutation sites, which is highlighted in red asterisk. The right panel shows the corresponding quantification of sanger sequencing.

(E, F) The knockout efficiency of NSUN2 (E) and NSUN6 (F) in HEK293T cell lines verified by Western blotting. The protein level of α-Tubulin and GAPDH were served as loading controls, separately.

(G, H) DRAM induced mutations close to m5C sites in AP5Z1 (G) and RPSA (H) mRNAs after NSUN2 and NSUN6 knockout in HEK293T cells. The left panel illustrates the location of DRAM induced mutation sites, which is highlighted in red asterisk. The right panel shows the corresponding quantification of sanger sequencing.

Earlier version: NSUN243 and NSUN644, two family members of NOL1/NSUN protein, were both identified as m5C methyltransferase of mRNA45. To verify that the detection of DRAM occurs in the presence of m5C, we individually depleted NSUN2 and NSUN6 in HEK293T cells and performed DRAM transfection. The knockout efficiency has been confirmed by western blotting (Fig.2E, 2F and Fig. S2, Supporting Information). It has been previously demonstrated that m5C methylation of AP5Z1 and RPSA are catalyzed by NSUN2 and NSUN6, respectively21,46. In line with this, sanger sequencing following RT-PCT showed a significant reduction in C-to-U or A-to-G mutations near the m5C sites in methyltransferase-deficient cells compared with WT cells (Fig. 2G and H). Overall, these findings suggest that the DRAM detection system is assayed in an m5C-dependent form.

Revised version: NSUN243 and NSUN644, two family members of NOL1/NSUN protein, were both identified as m5C methyltransferase of mRNA45. To verify that the detection of DRAM occurs in the presence of m5C, we individually depleted NSUN2 and NSUN6 in HEK293T cells and performed DRAM transfection. The knockout efficiency has been confirmed by western blotting (Fig.2E, 2F and Fig. S4A,4B). It has been previously demonstrated that m5C methylation of AP5Z1 and RPSA is catalyzed by NSUN2 and NSUN6, respectively21,46. In line with this, sanger sequencing following RT-PCT showed a significant reduction in C-to-U or A-to-G mutations near the m5C sites in methyltransferase-deficient cells compared with WT cells (Fig. 2G and H). Overall, these findings suggest that the DRAM detection system is assayed in an m5C-dependent form.

DRAM enables transcriptome-wide analysis of m5C methylation

Earlier version: Subsequently, we performed RNA-seq after DRAM transfection to enable a transcriptome-wide detection of m5C. To serve as positive controls, two previously published BS-seq datasets were also integrated5,21. Mutations were detected near the m5C site in RPSA as A-to-G by DRAM-ABE (Fig.3A), and DRAM-CBE detected the presence of C-to-U mutations near the AP5Z1 m5C site (Fig.3B). However, the DRAMmut and Deaminase systems induced few effective mutations close to these sites. Examination of multiple reported high-confidence RNA m5C sites showed that DRAM-seq editing events were also enriched in the vicinity of the BS-seq sites (Fig.3A, 3B and Fig. S3, Supporting Information). Comparison of three biological replicates revealed a strong reproducibility of A-to-G/C-to-U mutations in HEK293T cells expressing DRAM-ABE and DRAM-CBE (Fig. S4, Supporting Information). This finding suggests that DRAM selectively targets specific RNAs for editing, exhibiting a high degree of consistency across samples.

Revised version: Subsequently, we performed RNA-seq analysis after DRAM transfection by detecting C-to-U/A-to-G editing events to accomplish transcriptome-wide detection of m5C (Fig.3A). To serve as positive controls, two previously published BS-seq datasets were also integrated5,21. Mutations were detected near the m5C site in RPSA as A-to-G by DRAM-ABE (Fig.3B), and DRAM-CBE detected the presence of C-to-U mutations near the AP5Z1 m5C site (Fig.3C). However, the DRAMmut and Deaminase systems induced few effective mutations close to these sites. Examination of multiple reported high-confidence RNA m5C sites showed that DRAM-seq editing events were also enriched in the vicinity of the BS-seq sites (Fig.3B, 3C and Fig. S5).

DRAM enables transcriptome-wide analysis of m5C methylation.

Added in the revised version: (A) Schematic of the DRAM-seq method.

Earlier version: (A, B) Integrative genomics viewer (IGV) browser traces of DRAM-seq data expressing the indicated constructs in RPSA (A, left panel), TARBP2 (A, right panel), AP5Z1(B, left panel), and TRAF7 (B, left panel) mRNAs. C-to-U or A-to-G mutations found in at least 10% of reads are indicated by coloring. The previously published RNA BS-seq datasets from two individual studies were displayed as panel “Yang et al.” and “Zhang et al.”. (n(DRAM)=3 independent samples, n(Deaminase)=2 independent samples, and n(DRAMmut)=1 independent sample.)

Revised version: (B, C) Integrative genomics viewer (IGV) browser traces of DRAM-seq data expressing the indicated constructs in RPSA (B, left panel), TARBP2 (B, right panel), AP5Z1(C, left panel), and TRAF7 (C, left panel) mRNAs. C-to-U or A-to-G mutations found in at least 10% of reads are indicated by coloring. The previously published RNA BS-seq datasets from two individual studies were displayed as panel “Yang et al.” and “Zhang et al.”. (n(DRAM)=3 independent samples, n(Deaminase)=2 independent samples, and n(DRAMmut)=1 independent sample.)

Earlier version: (C, D) Integrative genomics viewer (IGV) browser traces of DRAM-seq data in wildtype and methyltranferases knockout cells in AP5Z1 (C) and RPSA (D) mRNAs. C-to-U or A-to-G mutations were found in at least 10% of reads are indicated by coloring. The previously published RNA BS-seq datasets from two individual studies were displayed as panel “Yang et al.” and “Zhang et al.”. n=3 independent samples.

Revised version: (D, E) Integrative genomics viewer (IGV) browser traces of DRAM-seq data in wildtype and methyltranferases knockout cells in AP5Z1 (D) and RPSA (E) mRNAs. C-to-U or A-to-G mutations were found in at least 10% of reads are indicated by coloring. The previously published RNA BS-seq datasets from two individual studies were displayed as panel “Yang et al.” and “Zhang et al.”. n=3 independent samples.

Added in the revised version: (F) Screening process for DRAM-seq assays and principles for screening high-confidence genes.

Earlier version: (E) The pie chart shows the distribution of editing sites in different transcript region in cells expressing DRAM (n=3 independent samples).

Revised version: (G) The pie chart shows the distribution of editing sites in different transcript region in cells expressing DRAM (n=3 independent samples).

Earlier version: (F) The density map showing the distribution of editing events across the mRNA transcripts detected by DRAM-seq.

Revised version: (H) The density map showing the distribution of editing events across the mRNA transcripts detected by DRAM-seq.

Earlier version: (G) The frequency plot shows the distribution of the distances of edit events in DRAM-seq relative to the m5C sites from the published BS-seq datasets. The position of each m5C site of BS-seq is determined as 0, and the relative distance of each site to the nearest edit event in DRAM-seq is calculated and plotted. The plots are presented separately based on the cutoff of upstream and downstream 3000bp (above) and 80bp (below) windows.

Revised version: (I) The frequency plot shows the distribution of the distances of edit events in DRAM-seq relative to the m5C sites from the published BS-seq datasets. The position of each m5C site of BS-seq is determined as 0, and the relative distance of each site to the nearest edit event in DRAM-seq is calculated and plotted. The plots are presented separately based on the cutoff of upstream and downstream 3000bp (above) and 80bp (below) windows.

Earlier version: (H, I) Motif analysis discovered within the ±20nt region around the C-to-U or A-to-G editing site in cells expressing DRAM-CBE (H), APOBEC1(H), DRAM-ABE (I) and TadA-8e (I).

Revised version: (J, K) Motif analysis discovered within the ±20nt region around the C-to-U or A-to-G editing site in cells expressing DRAM-CBE (J), APOBEC1(J), DRAM-ABE (K) and TadA-8e (K).

Earlier version: DRAM-seq analysis further confirmed that DRAM was detected in an m5C-dependent manner, with minimal mutations in AP5Z1 and RPSA mRNAs in methyltransferase knockout cells compared to wild-type cells (Fig. 3C, D). Moreover, the knockout cells exhibited overall rare DRAM-seq editing events close to m5C sites in other mRNAs (Fig. S5, Supporting Information).

Revised version: DRAM-seq analysis further confirmed that DRAM was detected in an m5C-dependent manner and that mutations in AP5Z1 and RPSA mRNA were reduced in methyltransferase knockout cells compared to wild-type cells (Fig. 3D, 3E). Moreover, the knockout cells exhibited overall rare DRAM-seq editing events close to m5C sites in other mRNAs (Fig. S6). Unfortunately however due to the inability of the DRAM detection system to achieve m5C detection with single base resolution, we did not find sequence preferences for the m5C sites responsible for catalysis by each of the two methyltransferases NSUN2 and NSUN6(Fig. S4D).

Added in the revised version: A comparison of three biological replicates from each experimental group revealed a strong reproducibility of A-to-G/C-to-U mutations in HEK293T cells expressing DRAM-ABE and DRAM-CBE (Fig. S7). Moreover, the integration and analysis of the experimental data revealed a high degree of overlap between the three biological replicates (Fig. S4C). And a recent study by Wang et al. showed that ALYREF deletion affects the expression of 94 mRNAs47, and only 55.32% of these ALYREF-regulated mRNAs can be detected by the DRAM system (Fig. S4E). These findings suggest that DRAM selectively targets specific RNAs for editing, exhibiting a high degree of consistency across samples.

Earlier version: To obtain information on a set of high-confidence DRAM-seq data, we filtered the list of sites transfected with deaminase alone and screened the sequencing results with methyltransferase depleted, pooled editing events occurring in at least 10% of reads across multiple samples to obtain a set of high-confidence editing sites (Supplementary Table 2), and integrated genes with editing sites occurring in DRAM-ABE and DRAM-CBE (Supplementary Table 3).

Revised version: To obtain information on a set of high-confidence DRAM-seq data, we filtered the list of sites transfected with deaminase alone and screened the sequencing results with methyltransferase depleted, pooled editing events occurring in at least 10% of reads across multiple samples to obtain a set of high-confidence editing sites (Fig. 3F and Supplementary Table 2), and integrated genes with editing sites occurring in DRAM-ABE and DRAM-CBE (Fig. 3F and Supplementary Table 3).

Earlier version: Previous studies have indicated that m5C sites are predominantly distributed in the coding sequences (CDS) and notably enriched near the initiation codon5,25,26,47-49. To further delineate the characteristics of the DRAM-seq data, we compared the distribution of DRAM-seq editing sites within the gene structure, specifically examining their occurrences in the 5’untranslated region (5’UTR), 3’ untranslated region (3’UTR), CDS and ncRNA. Our analysis revealed that DRAM-seq editing events in cells expressing DRAM-ABE and DRAM-CBE were primarily located in the CDS and 3’UTR, indicating a non-random distribution of m5C (Fig.3E, Fig. S6A and 6B, Supporting Information). Moreover, plotting the distribution of DRAM-seq editing sites in mRNA segments (5’UTR, CDS, and 3’UTR) highlighted a significant enrichment near the initiation codon (Fig.3F). In contrast, cells expressing the deaminase exhibited a distinct distribution pattern of editing sites, characterized by a prevalence throughout the 5’UTR. This finding reaffirms that the specific editing pattern observed in DRAM-seq across the transcriptome depends on its capacity to bind m5C (Fig.3F).

Revised version: Previous studies have indicated that m5C sites are predominantly distributed in the coding sequences (CDS) and notably enriched near the initiation codon5,25,26,48–50. To further delineate the characteristics of the DRAM-seq data, we compared the distribution of DRAM-seq editing sites within the gene structure, specifically examining their occurrences in the 5’untranslated region (5’UTR), 3’ untranslated region (3’UTR), CDS and Intergenic/Intron region. Our analysis revealed that DRAM-seq editing events in cells expressing DRAM-ABE and DRAM-CBE were primarily located in the CDS and 3’UTR, indicating a non-random distribution of m5C (Fig.3G, Fig. S8A and 8B). Moreover, plotting the distribution of DRAM-seq editing sites in mRNA segments (5’UTR, CDS, and 3’UTR) highlighted a significant enrichment in the CDS (Fig.3H). In contrast, cells expressing the deaminase exhibited a distinct distribution pattern of editing sites, characterized by a prevalence throughout the 3’UTR (Fig.3H). This finding reaffirms that the specific editing pattern observed in DRAM-seq across the transcriptome depends on its capacity to bind m5C.

Earlier version: Comparative analysis of the DRAM-seq editing sites with the previously published BS-seq m5C sites indicated that the likelihood of editing was notably higher in closer proximity to the m5C sites (Fig.3G). Furthermore, the editing window of DRAM exhibited enrichment approximately 20bp before and after the m5C site (Fig.3G). Investigation into the sequences surrounding the editing window revealed that AC motifs were the most significantly enriched in DRAM-CBE, whereas (U/C) A motifs were most notably enriched in DRAM-ABE. In contrast, the APOBEC1 and TadA-8e samples displayed no significantly enriched motifs, with mutations being more randomly orientated (Fig.3H, I).

Revised version: Comparative analysis of the DRAM-seq editing sites with the previously published BS-seq m5C sites indicated that the likelihood of editing was notably higher in closer proximity to the m5C sites (Fig.3I). Furthermore, the editing window of DRAM exhibited enrichment approximately 20bp before and after the m5C site (Fig.3I). Investigation into the sequences surrounding the editing window revealed that AC motifs were the most significantly enriched in DRAM-CBE, whereas (U/C) A motifs were most notably enriched in DRAM-ABE. In contrast, the APOBEC1 and TadA-8e samples displayed no significantly enriched motifs, with mutations being more randomly orientated (Fig.3J, 3K).

DRAM-seq provides stable and comprehensive identification of m5C loci

Earlier version: Subsequently, we assessed the capacity of DRAM-seq to detect m5C on a transcriptome-wide scale, comparing its performance to BS-seq that have been previously reported with great authority. Although both previous studies employed bisulfite treatment, the resulting data obtained significantly discrepancies due to variations in their treatment and analysis methodologies. Our results indicated that DRAM-seq identified the presence of m5C modifications covering 82.5% of the genes detected by Yang et al.5 and 95.5% of the genes detected by Zhang et al.21 (Fig.4A and C). Remarkably, certain pivotal regulators with diverse biological functions, such as ATG16L1(coordinats autophagy pathway)50 and ARHGEF25 (plays an important role in actin cytoskeleton reorganisation)51, were identified by Zhang et al. and DRAM-seq, but not by Yang et al. (Fig.4B). Conversely, FANCD2 (Maintains chromosome stability)52 and RPL15(components of the large ribosomal subunit)53,54, were discovered by Yang et al. and DRAM-seq, but not by Zhang et al. (Fig.4D). Hence, DRAM-seq appears to offer a more stable and comprehensive identification of the m5C loci.

Revised version: Subsequently, we then evaluated the ability of DRAM-seq to detect m5C across the entire transcriptome and compared its performance to that of the previously reported BS-seq. Although both previous studies employed bisulfite treatment, the resulting data obtained significant discrepancies due to variations in their treatment and analysis methodologies. Our results indicated that DRAM-seq identified the presence of m5C modifications covering 79.6% of the genes detected by Yang et al.5 and 91.9% of the genes detected by Zhang et al.21 (Fig.4A and C). Remarkably, certain pivotal regulators with diverse biological functions, such as ATG16L1(coordinats autophagy pathway)51 and ARHGEF25 (plays an important role in actin cytoskeleton reorganisation)52, were identified by Zhang et al. and DRAM-seq, but not by Yang et al. (Fig.4B). Conversely, FANCD2 (Maintains chromosome stability)53 and RPL15(components of the large ribosomal subunit)54,55, were discovered by Yang et al. and DRAM-seq, but not by Zhang et al. (Fig.4D). Hence, DRAM-seq appears to offer a more stable and comprehensive identification of the m5C loci.

DRAM monitors m5C dynamics in cellular RNA.

(A) Venn diagram showing the overlap between DRAM-seq and Yang et al.’s edited genes.

(C) Venn diagram showing the overlap between DRAM-seq and Zhang et al.’s edited genes.

(B, D) Integrative genomics viewer (IGV) browser traces of DRAM-seq data expressing the indicated constructs in the ATG16L1(B), ARHGEF25(B), FANCD2(D), and RPL15(D) mRNAs. C-to-U/A-to-G mutations found in at least 10% of reads are indicated by coloring, and the m5C site found by BS-seq is also labelled.

(E) Genes with DRAM-seq editing events were analyzed for KEGG bioprocess enrichment.

(F) GO biological processes enrichment analysis of genes with DRAM-seq editing events. Statistical analyses were performed using the DAVID tool.

Earlier version: To provide functional insights into m5C RNA-modified genes in HEK293T cells, we conducted Gene Ontology (GO) and Kyoto Encyclopedia of Genes and Genomes (KEGG) analyses. These results highlighted the involvement of these genes in the regulation of diverse key biological processes, such as cell division, cell cycle, mRNA splicing, protein processing in the endoplasmic reticulum, nucleocytoplasmic transport, translation, DNA repair and others (Fig.4E, 4F, Fig S6C and S6D, Supporting Information).

Revised version: To provide functional insights into m5C RNA-modified genes in HEK293T cells, we conducted Gene Ontology (GO) and Kyoto Encyclopedia of Genes and Genomes (KEGG) analyses. These results highlighted the involvement of these genes in the regulation of diverse key biological processes, such as cell division, cell cycle, mRNA splicing, protein processing in the endoplasmic reticulum, nucleocytoplasmic transport, translation, DNA repair and others (Fig.4E, 4F, Fig. S8C and S8D).

DRAM monitors m5C dynamics in cellular RNA

Earlier version: Next, we aim to determine whether changes in m5C can be detected in real-time by DRAM. Previous studies have established that oxidative stress induced by 200 μM sodium arsenite (NaAsO2) treatment inhibited the expression of NSUN2 and NSUN6, resulting in a gradual decrease in m5C methylation55,56. As cells exposed to NaAsO2 typically exhibited slowed cell cycle progression and NaAsO2 induce apoptosis, we performed a low concentration screen57,58. After gradient treatment of NaAsO2 (50μM, 100μM, 150μM, and 200μM, separately) in cells, 200μM group was identified as the optimal concentration, leading to highly significant decreases of NSUN2 and NSUN6 expression (Fig.5A). Subsequently, HEK293T cells expressing DRAM with NaAsO2 for 3 hours revealed a decreasing trend in the mRNA expression of NSUN2 and NSUN6 (Fig.5B).

Revised version: Next, we aim to determine whether changes in m5C can be detected in real time by DRAM. Previous studies have established that oxidative stress induced by 200 μM sodium arsenite (NaAsO2) treatment inhibited the expression of NSUN2 and NSUN6, resulting in a gradual decrease in m5C methylation56,57. As cells exposed to NaAsO2 typically exhibited slowed cell cycle progression and NaAsO2-induced apoptosis, we performed a low concentration screen58,59. After gradient treatment of NaAsO2 (50μM, 100μM, 150μM, and 200μM, separately) in cells, 200μM group was identified as the optimal concentration, leading to highly significant decreases of NSUN2 and NSUN6 expression (Fig.5A). Subsequently, HEK293T cells expressing DRAM with NaAsO2 for 3 hours revealed a decreasing trend in the mRNA expression of NSUN2 and NSUN6 (Fig.5B). Bisulfite sequencing PCR also verified a decreasing trend in m5C levels in selected mRNAs after NaAsO2 treatment (Fig. 5C and D).

DRAM monitors m5C dynamics in cellular RNA.

(A) Relative expression levels of NSUN2 and NSUN6 mRNA in HEK293T cells treated with 50 μM, 100 μM, 150μM and 200μM NaAsO2.

(B) Relative expression levels of NSUN2 and NSUN6 mRNA in DRAM-expressing HEK293T cells treated with 200μM NaAsO2 for 3 hours.

Removed in the revised version: (C, D) The m5C adjacent mutations from RPSA (C) and AP5Z1 (D) mRNA in DRAM-expressing HEK293T cells treated with NaAsO2. The left panel indicates the sanger sequencing results following RT-PCR, while the corresponding quantifications of DRAM-induced mutations are shown on the right panel.

Earlier version: (E, F) BS-PCR showing the m5C levels in RPSA and AP5Z1 mRNAs after 3 h of control and NaAsO2 treatment in cells.

Revised version: (C, D) Bisulfite sequencing PCR showing the m5C levels in RPSA and AP5Z1 mRNAs after 3 h of control and NaAsO2 treatment in cells.

Added in the revised version: (E, F) The m5C adjacent mutations from RPSA (C) and AP5Z1 (D) mRNA in DRAM-expressing HEK293T cells treated with NaAsO2. The left panel indicates the sanger sequencing results following RT-PCR, while the corresponding quantifications of DRAM-induced mutations are shown on the right panel.

(G, H) DRAM analysis of RPSA (G) and AP5Z1 (H) mRNAs with 250 ng, 50 ng, and 10 ng of input RNA. Representative Sanger sequencing plots are shown on the left panel, with mutation sites marked with asterisks. The mutation rates are quantified on the right panel.

(I) Flowchart illustrating Cell viability analysis by CCK8 reagent after DRAM transfection in HEK293T cells.

(J) Quantitative comparison of the relative proliferative capacity of DRAM-expressing and untransfected cells.

Earlier version: Sanger sequencing results showed that induction of the DRAM system after NaAsO2 treatment resulted in a highly significant reduction of A-to-G/C-to-U mutations on both RPSA and AP5Z1 mRNAs compared to cells without NaAsO2 treatment (Fig.5C and D). To confirm this analysis, we employed bisulfite sequencing PCR to verify that m5C levels in selected mRNAs displayed a decreasing trend after NaAsO2 treatment (Fig.5E and F). Thus, the DRAM effectively detects changes in m5C, providing a valuable tool for monitoring the accumulation of m5C in individual RNAs in response to alterations cellular conditions.

Revised version: Next, we performed DRAM system assays, and Sanger sequencing results showed that induction of the DRAM system after NaAsO2 treatment resulted in a highly significant reduction of A-to-G/C-to-U mutations on both RPSA and AP5Z1 mRNAs compared to cells without NaAsO2 treatment (Fig.5E and F). Taken together, with our previous evidence that the methyltransferase-deleted cells exhibited overall rare DRAM editing events close to m5C sites, we suggest that the DRAM could effectively monitor the accumulation of m5C in individual RNAs in response to the alteration of cellular conditions.

DRAM enables low-input m5C profiling

Earlier version: A significant challenge in m5C detection lies in the specificity of antibodies and the substantial amount of input RNA required for sequencing. RNA is susceptible to degradation during denaturation, sodium bisulfite treatment and desulfurization steps in the BS-seq assay59. Immunoprecipitation-based m5C assays and LC-MS/MS also impose high demand for sample input7,25,60. Several experiments have highlighted the requirement of 100-500 ng of RNA for m5C-RIP-seq, while BS-seq necessitates an even more demanding 500-750μg of RNA21,25,61. To assess the detection limits of DRAM-Sanger, we attempted to amplify two representative m5C-containing sites in the RPSA and AP5Z1 transcripts from diluted RNA samples.

Revised version: A significant challenge in m5C detection lies in the specificity of antibodies and the substantial amount of input RNA required for sequencing. RNA is susceptible to degradation during denaturation, sodium bisulfite treatment and desulfurization steps in the BS-seq assay60. Immunoprecipitation-based m5C assays and LC-MS/MS also impose high demand for sample input7,25,61. Several experiments have highlighted the requirement of 100-500 ng of RNA for m5C-RIP-seq, while BS-seq necessitates an even more demanding 750-1000 ng of RNA21,25,62. To assess the detection limits of DRAM-Sanger, we attempted to amplify two representative m5C-containing sites in the RPSA and AP5Z1 transcripts from diluted RNA samples.

Remarkably, we successfully generated PCR products of these two mRNAs from cDNAs corresponding to 250 ng, 50 ng, and 10 ng of total RNA. Quantitative analysis by Sanger sequencing demonstrated nearly identical Sanger traces across these dilutions (Fig.5G and H). This finding underscores that the specificity of DRAM editing depended on its ability to bind m5C, and DRAM is proficient in low-input m5C analyses. Furthermore, cell viability was determined by CCK8 assay on HEK293T cells transfected with DRAM (Fig.5I). Importantly, there was no significant difference in the relative proliferative capacity of the cells compared to untransfected cells (NC), indicating that DRAM expression did not adversely affect cell viability (Fig.5J).

Added in the revised version: Transfection of the DRAM system in cells results in the transient overexpression of fusion proteins. To investigate how varying expression levels of these proteins influence A-to-G and C-to-U editing within the same m5C region, we conducted a gradient transfection using plasmid concentrations of 1500 ng, 750 ng and 300 ng. This approach allowed us to progressively reduce the expression levels of the fusion proteins (Fig. S9A and S9B). Sanger sequencing revealed that the editing efficiency of A-to-G and C-to-U within the m5C region significantly decreased as fusion protein expression diminished (Fig. S9C and S9D). These findings suggest that the transfection efficiency of the DRAM system is concentration-dependent and that the ratio of editing efficiency to transfection efficiency may assist in the quantitative analysis of m5C using the DRAM system.

Removed in the revised version: Conclusions

Removed in the revised version: In summary, we developed a novel deaminase and reader protein assisted RNA m5C methylation approach, which does not rely on antibodies or bisulfite, thus leading to unprecedently comprehensive transcriptome-wide RNA m5C methylation profiling. We anticipated that this system could pave the way for uncovering further biological functions of m5C modifications and facilitate the development of therapeutic interventions for associated diseases.

Discussion

In recent years, m5C methylation modifications have received increasing attention, with multiple reports detailing the distribution of RNA m5C methylation modifications across various species and tissues, elucidating their characteristics. Despite the relatively low abundance of m5C, its highly dynamic changes hold significant implications for the regulation of physiological and pathological processes5,21,44. However, due to the limitations of sequencing methods and the variability of data processing, there remains ample room for progress in the study of m5C detection methods.

Earlier version: In this study, we developed site-specific, depth-sequencing-free m5C detection method using DRAM-Sanger. This workflow relies on conventional molecular biology assays such as RT-PCR and Sanger sequencing, eliminating the need for specialized techniques and thereby simplifying the process of m5C detection.

Revised version: In this study, we developed a site-specific, depth-sequencing-free m5C detection method using DRAM-Sanger. This workflow relies on conventional molecular biology assays such as RT-PCR and Sanger sequencing, eliminating the need for specialized techniques and thereby simplifying the process of m5C detection.

Earlier version: DRAM-seq introduce a novel strategy for transcriptome-wide m5C detection, overcoming inherent limitations in existing methods. Notably, DRAM-seq covered 82% of the high-confidence m5C modified genes detected by BS-seq, and identified more potential m5C sites. This can be attribute to the avoidance of bisulfite treatment by DRAM-seq, preventing RNA damage and ensuring a more comprehensive representation of RNA samples. This feature also likely contributes to the observed stability of DRAM-seq in comparison to BS-seq. Additionally, DRAM-seq is not constrained by antibody specificity and is resistant to chemical reagent-induced damage, preserving the gene expression information captured by typical RNA-seq. Consequently, DRAM-seq uniquely provides the capability to report both transcript abundance and information related to m5C modifications.

Revised version: DRAM-seq introduces a novel strategy for transcriptome-wide m5C detection, overcoming inherent limitations in existing methods. Notably, DRAM-seq covered around 80% of the high-confidence m5C-modified genes detected by BS-seq and identified more potential m5C sites. This can be attributed to the avoidance of bisulfite treatment by DRAM-seq, preventing RNA damage and ensuring a more comprehensive representation of RNA samples. This feature also likely contributes to the observed stability of DRAM-seq in comparison to BS-seq. Additionally, DRAM-seq is not limited by antibody specificity and is resistant to chemical-induced damage.

A prominent challenge in existing m5C profiling methods is their reliance on substantial amounts of input RNA samples. In contrast, DRAM operates through the deamination activity of deamination activity of deaminase, preserving RNA integrity and preventing degradation. The notable advantage of DRAM lies in its capacity for low-input m5C detection. Our analysis demonstrates that DRAM requires as low as 10ng of total RNA for m5C detection. While DRAM is currently well-suited for detecting m5C on a transcriptome-wide scale, the potential for future applications involving third-generation sequencing could extend its utility to individual mRNAs, particularly m5C heterogeneity on mRNA splicing variants. In addition, the DRAM system depends on the specific recognition of m5C modifications on ssRNA by the reader protein, theoretically avoiding the false-positive effects of 5-hydroxymethylation modifications in other assays, such as BS-seq21–23. This potential feature could enhance the accuracy of the DRAM assay, albeit it still requires careful validation.

Earlier version: In our study, m5C detection was performed following the transient transfection of the DRAM detection system into mammalian cells, which might result in a lower mutation rate at the corresponding site. Therefore, employing lentiviral-mediated transfection into cell lines of interest could potentially enhance the efficiency of m5C detection. Further elucidation of the key amino acids directing ALYREF and YBX1’s binding to m5C methylation sites should enable more accurate and sensitive m5C detection by DRAM-seq. Since m5C lacks a fixed motif, DRAM has an apparent limitation in achieving single-base resolution for detecting m5C. However, our present study proved that the measuring resolution of DRAM is around 40nt, which facilitates higher precision than that of m5C-RIP-seq (∼100nt). In the future, with the in-depth structural analysis of the m5C current readers and the identification of new m5C readers, we anticipated to better determine the motif of m5C for its accurate localization. Moreover, the substitution of deaminases, such as A3A and A3G (the family members of APOBEC), could also potentially enhance the efficiency of the DRAM detection 62-64.

Revised version: In our study, m5C detection was performed following the transient transfection of the DRAM detection system into mammalian cells, which might result in a lower mutation rate at the corresponding site. Therefore, employing lentiviral-mediated transfection into cell lines of interest could potentially enhance the efficiency of m5C detection. In the present study, the editing efficiency of the DRAM system, which is believed to detect dynamic changes in m5C, was significantly reduced following sodium arsenite treatment. However, sodium arsenite can adversely affect all proteins containing thiol groups, impacting not only NSUN proteins but potentially also the function of reader proteins and deaminases. As a result, this validation cannot entirely rule out the possibility that the observed decrease in editing frequency is due to the inactivation of fusion proteins by sodium arsenite. The optimization of this validation may require the development of more specific m5C inhibitors in the future. Further elucidation of the key amino acids directing ALYREF and YBX1’s binding to m5C methylation sites should enable more accurate and sensitive m5C detection by DRAM-seq. Due to the lack of a fixed base composition for characterizing all m5C modification sites, DRAM has an apparent limitation in achieving single-base resolution for detecting m5C. However, our present study proved that the measuring resolution of DRAM is around 40nt, which facilitates higher precision than that of m5C-RIP-seq (∼100nt). In the future, with more in-depth analyses of m5C reader structures and the identification of new potential m5C readers, we expect to achieve more precise m5C localization and more comprehensive m5C modification detection. Moreover, the substitution of deaminases, such as A3A and A3G (the family members of APOBEC), could also potentially enhance the efficiency of the DRAM detection 63–65.

Earlier version: Following the purification of DRAM fusion protein, achieving in vitro detection of RNA m5C methylation could broaden the applicability to a more diverse range of samples. Another potential application for DRAM-seq could be the expression of drug-inducible DRAM systems in vivo using various animal models for m5C analysis. These will together provide novel insights of m5C modifications for biological and clinical research.

Revised version: Following the purification of DRAM fusion protein, achieving in vitro detection of RNA m5C methylation could broaden the applicability to a more diverse range of samples. Another potential application for DRAM-seq could be the expression of drug-inducible DRAM systems in vivo using various animal models for m5C analysis. These will together provide novel insights into m5C modifications for biological and clinical research.

Added in the revised version: Conclusions

Added in the revised version: In summary, we developed a novel deaminase and reader protein-assisted RNA m5C methylation approach that detects the m5C region by deaminating As or Cs in close proximity to the m5C sites, which does not rely on antibodies or bisulfite, thus leading to unprecedently comprehensive transcriptome-wide RNA m5C methylation profiling. We anticipated that this system could pave the way for uncovering further biological functions of m5C modifications and facilitate the development of therapeutic interventions for associated diseases.

Materials and methods

Plasmid construction

ALYREF and YBX1 expression plasmids were purchased from MIAOLING BIOLOGY (http://www.miaolingbio.com/), and the ALYREF and YBX1 fractions were then amplified using specific primer. The ALYREF and YBX1 portions were amplified using pCMV-APOBEC1-YTH (Addgene plasmid no. 131636; https://www.addgene.org/131636/) and ABE8e (Addgene plasmid no. 138489; https://www.addgene.org/ 138489/) to amplify the deaminase portion and the essential plasmid construct proxies, and finally the fragments were recombined by the ClonExpress Ultra One Step Cloning Kit to complete the plasmid vector construction. Both DRAMmut-ABE and DRAMmut-CBE related vectors were obtained by introducing the corresponding key amino acid mutations using Fast Site-Directed Mutagenesis Kit (TIANGEN Biotech). The primer sequences used are listed in Supplementary Table 4.

Cell culture and plasmid transfection

HEK293T cell line (ATCC) was cultured in Dulbecco’s Modified Eagle Medium (DMEM) supplemented with 10% fetal bovine serum (CLARK BIOSCIENCE) and 1% penicillin (100 U/ml)-streptomycin (100μg/ml). The cells were seeded in 12-well plates and transfected using Hieff Trans™ Liposomal Transfection Reagent (Yeasen).

NSUN2-depleted cell lines were generated by cloning NSUN2-targeting single guide RNA sequences into the pSpCas9(BB)−2A-Puro (PX459) V2.0 plasmid (Addgene plasmid no. 62988; http://n2t.net/addgene:62988). Plasmids were then transfected into HEK293T cells and Puromycin (Meilunbio) was added at a final concentration of 3 μg/ml to enrich the positively transfected cells 24 h after transfection. After 72 h, the cells were collected and used for genotyping by Sanger sequencing. NSUN6-depleted cell lines were generated in the same way. The primers used for genotyping and single guide RNA sequences are listed in Supplemental Table 4.

Cell viability measurements

HEK293T cells were transfected with DRAM plasmid and cultured at 37°C for 24 h. Subsequently, 1000 cells were seeded in 96-well plates. After waiting for the cells to attach to the wall, the cell activity was detected by Cell Counting Kit-8 (Meilunbio). Cell Counting Kit-8 contains WST-8, which in the presence of the electronically coupled reagent 1-Methoxy PMS can be reduced by mitochondrial dehydrogenase to the orange-colored metazan product Formazan, the absorbance of which is measured at 450 nm to analyze cellular activity.

Western blotting

For protein blotting, samples were lysed in RIPA Lysis Buffer (Meilunbio) with Phenylmethanesulfonyl fluoride (PMSF) and the BCA protein assay kit (Beyotime Biotechnology) was used to Protein concentration was measured. Total protein extracts were separated by SDS-PAGE on a 10% gel and then transferred to 0.22 nm polyvinylidene fluoride membranes (Boster). Subsequently, the proteins were probed with specific antibodies after the blot was blocked with 5% non-fat milk (Boster). Images were quantified using ImageJ software and all data are expressed as mean ± SEM.

The following antibodies and concentrations were used: NSUN2 Polyclonal antibody (Proteintech; Cat No.20854-1-AP; 1:7500), NSUN6 Polyclonal antibody (Proteintech; Cat No. 17240-1-AP; 1:2000), RabbitAnti-GAPDH antibody (Bioss; bs-41373R; 1:2000), Alpha Tubulin Polyclonal antibody (Proteintech; Cat No. 11224-1-AP; 1:2000), HRP-labeled Goat Anti-Rabbit IgG(H+L) (Beyotime Biotechnology; A0208; 1:2000).

cDNA synthesis and Sanger sequencing

Earlier version: Total cellular RNA was extracted with TRIzol reagent (TIANGEN Biotech) and cDNA was synthesized using PrimeScript™ II 1st Strand cDNA Synthesis Kit (Takara Bio) according to the manufacturer’s recommendations. PCR was then performed using 2×Taq PCR MasterMix II (TIANGEN Biotech) and primers flanking m5C target sites, and the purified PCR products were directly sequenced by Sanger sequencing. The Sanger sequencing results were analyzed using EditR 1.0.10 to calculate the mutation frequency65. The primers used in this study are shown in Supplemental Table 4.

Revised version: Total cellular RNA was extracted with TRIzol reagent (TIANGEN Biotech) and cDNA was synthesized using PrimeScript™ II 1st Strand cDNA Synthesis Kit (Takara Bio) according to the manufacturer’s recommendations. PCR was then performed using 2×Taq PCR MasterMix II (TIANGEN Biotech) and primers flanking m5C target sites, and the purified PCR products were directly sequenced by Sanger sequencing. The Sanger sequencing results were analyzed using EditR 1.0.10 to calculate the mutation frequency66. The primers used in this study are shown in Supplemental Table 4.

Real-time quantitative PCR

cDNA was synthesized using FastKing RT kit (with gDNase) (TIANGEN Biotech) according to the manufacturer’s recommendations. RT-qPCR assay was performed using SuperReal PreMix Plus (SYBR Green) (TIANGEN Biotech). GAPDH was used as an endogenous control, and the expression levels were normalized to the control and calculated by the 2-ΔΔCt formula. All samples were analyzed in triplicate and each mRNA quantification represents the average of at least three measurements. All data are expressed as mean ± SEM. The primers used in this study are shown in Supplemental Table 4.

Protein structure modelling

Earlier version: Protein structure simulations were performed using the SWISS-MODEL online website (https://swissmodel.expasy.org/interactive)66. The SWISS-MODEL database is able to provide up-to-date annotated 3D protein models, which are generated from automated homology modelling of related model organisms and experimental structural information for all sequences in UniProtKB, with reliable structural information, and subsequently protein structure observations were performed using PyMOL67.

Revised version: Protein structure simulations were performed using the SWISS-MODEL online website (https://swissmodel.expasy.org/interactive)67. The SWISS-MODEL database is able to provide up-to-date annotated 3D protein models, which are generated from automated homology modelling of related model organisms and experimental structural information for all sequences in UniProtKB, with reliable structural information, and subsequently protein structure observations were performed using PyMOL68.

Bisulfite sequencing PCR

We referenced bisulfite sequencing PCR, an assay established by Matthias Schaefer et al. We chemically deaminated cytosine in RNA using the EZ RNA methylation kit (50) (ZYMO RESEARCH) and then quantified m5C methylation levels based on PCR amplification of cDNA combined with deep sequencing23.

Earlier version: RNA Conversion Reagent was premixed with prepared RNA samples, and the RNA was denatured at 70°C for 5 minutes, followed by a reaction period of 45 minutes at 54°C. Finally, the purified RNA samples were recovered after desulfurization by RNA Desulphonation Buffe. cDNA was synthesized using PrimeScript™ II 1st Strand cDNA Synthesis Kit (Takara Bio) according to the manufacturer’s recommendations. PCR was then performed using 2× EpiArt HS Taq Master Mix (Dye Plus) (Vazyme) and m5C target site specific Bisulfite Primer (primer sequences were designed at https://zymoresearch.eu/pages/bisulfite-primer-seeker), the products were purified by TIANgel Midi Purification Kit (TIANGEN Biotech), and the connectors for second-generation sequencing were attached at both ends of the products for sequencing. Finally, deep sequencing was performed by HiTOM analysis to detect the methylation level (The number of reads >1000 in deep sequencing) 68. The primers used in this study are shown in Supplemental Table 4.

Revised version: RNA Conversion Reagent was premixed with prepared RNA samples, and the RNA was denatured at 70°C for 5 minutes, followed by a reaction period of 45 minutes at 54°C. Finally, the purified RNA samples were recovered after desulfurization by RNA Desulphonation Buffe. cDNA was synthesized using PrimeScript™ II 1st Strand cDNA Synthesis Kit (Takara Bio) according to the manufacturer’s recommendations. PCR was then performed using 2× EpiArt HS Taq Master Mix (Dye Plus) (Vazyme) and m5C target site-specific Bisulfite Primer (primer sequences were designed at https://zymoresearch.eu/pages/bisulfite-primer-seeker), the products were purified by TIANgel Midi Purification Kit (TIANGEN Biotech), and the connectors for second-generation sequencing were attached at both ends of the products for sequencing. Finally, deep sequencing was performed by HiTOM analysis to detect the methylation level (The number of reads >1000 in deep sequencing) 69. The primers used in this study are shown in Supplemental Table 4.

Library construction and next-generation sequencing

Earlier version: 1μl of total cellular RNA was used for sequencing library generation by NEBNext Ultra RNA Library Prep Kit for Illumina (NEB, USA, Catalog #: E7530L) following manufacturer’s recommendations and index codes were added to attribute sequences to each sample. Briefly, mRNA was purified from total RNA using poly-T oligo-attached magnetic beads. Fragmentation was carried out using divalent cations under elevated temperature in NEB Next First Strand Synthesis Reaction Buffer(5X). First strand cDNA was synthesized using random hexamer primer and M-MuLV Reverse Transcriptase (RNase H). Second strand cDNA synthesis was subsequently performed using DNA Polymerase I and RNase H. Remaining overhangs were converted into blunt ends via exonuclease/polymerase activities. After adenylation of 3’ ends of DNA fragments, NEB Next Adaptor with hairpin loop structure were ligated to prepare for hybridization. To select cDNA fragments of preferentially 370∼420 bp in length, the library fragments were purified with AMPure XP system (Beverly, USA). Then 3 µL USER Enzyme (NEB, USA) was used with size selected, adaptor-ligated cDNA at 37°C for 15 min followed by 5 min at 95 °C before PCR. Then PCR was performed with Phusion High-Fidelity DNA polymerase, Universal PCR primers and Index (X) Primer. At last, PCR products were purified (AMPure XP system) and library quality was assessed on the Agilent 5400 system(Agilent,USA) and quantified by QPCR (1.5 nM). The qualified libraries were pooled and sequenced on Illumina platforms with PE150 strategy in Novogene Bioinformatics Technology Co., Ltd (Beijing, China), according to effective library concentration and data amount required.

Revised version: 1μg of total cellular RNA was used for sequencing library generation by NEBNext Ultra RNA Library Prep Kit for Illumina (NEB, USA, Catalog #: E7530L) following manufacturer’s recommendations and index codes were added to attribute sequences to each sample. Briefly, mRNA was purified from total RNA using poly-T oligo-attached magnetic beads. Fragmentation was carried out using divalent cations under elevated temperature in NEB Next First Strand Synthesis Reaction Buffer(5X). First-strand cDNA was synthesized using random hexamer primer and M-MuLV Reverse Transcriptase (RNase H). Second-strand cDNA synthesis was subsequently performed using DNA Polymerase I and RNase H. Remaining overhangs was converted into blunt ends via exonuclease/polymerase activities. After adenylation of 3’ ends of DNA fragments, NEB Next Adaptor with hairpin loop structure was ligated to prepare for hybridization. To select cDNA fragments of preferentially 370∼420 bp in length, the library fragments were purified with AMPure XP system (Beverly, USA). Then 3 µL USER Enzyme (NEB, USA) was used with size selected, adaptor-ligated cDNA at 37°C for 15 min followed by 5 min at 95 °C before PCR. Then PCR was performed with Phusion High-Fidelity DNA polymerase, Universal PCR primers and Index (X) Primer. At last, PCR products were purified (AMPure XP system) and library quality was assessed on the Agilent 5400 system(Agilent, USA)and quantified by QPCR (library concentration ≥ 1.5 nM). The qualified libraries were pooled and sequenced on Illumina platforms with PE150 strategy in Novogene Bioinformatics Technology Co., Ltd (Beijing, China), according to effective library concentration and data amount required.

Earlier version: DRAM-seq analysis and calling of edited cites

Revised version: DRAM-seq analysis and calling of edited sites

Earlier version: The raw fastq sequencing data were cleaned by trimming the adapter sequences using Fastp (v0.23.1) and were aligned to the human genome (hg19) using STAR (v2.7.7) in paired-end mode. The aligned BAM files were sorted and PCR duplicates were removed using Samtools (v1.12). The cite calling pf DRAM-seq was performed using Bullseye, a previously customized pipeline to look for C-to-U or A-to-G edited sites throughout the transcriptome40. Briefly, the sorted and deduplicated BAM files were initially parsed by parseBAM.pl script. Then, Find_edit_site.pl script was employed to find C-to-U or A-to-G editing events by DRAM-seq with at least 10 reads of coverage, an edit ratio of 5%-95%, an edit ratio at least 1.5-fold higher than NSUN2 or NSUN6-knockout samples, and at least 2 editing events at a given site. Sites that were only found in one or two replicates of each DRAM protein variant were removed. Editing events appeared in cells expressing merely APOBEC1 or TadA8e were also removed.

Revised version: The raw fastq sequencing data were cleaned by trimming the adapter sequences using Fastp (v0.23.1) and were aligned to the human genome (hg19) using STAR (v2.7.7) in paired-end mode. The aligned BAM files were sorted and PCR duplicates were removed using Samtools (v1.12). The cite calling pf DRAM-seq was performed using Bullseye, a previously customized pipeline to look for C-to-U or A-to-G edited sites throughout the transcriptome40. Briefly, the sorted and deduplicated BAM files were initially parsed by parseBAM.pl script.

Added in the revised version: Then, Find_edit_site.pl script was employed to find C-to-U or A-to-G editing events by DRAM-seq with at least 10 reads of coverage, an edit ratio of 5%-95%, an edit ratio at least 1.5-fold higher than NSUN2 or NSUN6-knockout samples, and at least 2 editing events at a given site. Sites that were only found in one replicate of each DRAM protein variant were removed. Editing events appeared in cells expressing merely APOBEC1 or TadA8e were also removed. For high confidence filtering, we further adjusted the Find_edit_site.pl parameters to the edit ratio of 10%-60%, an edit ratio of control samples at least 2-fold higher than NSUN2 or NSUN6-knockout samples, and at least 4 editing events at a given site.

Metagene and motif analyses

Earlier version: Metagene analysis was performed using hg19 annotations according to previously reported tool, MetaplotR69. For motif analysis, the 20bp flanking sequence of each DRAM-seq editing site were extracted by Bedtools (v2.30.0)70. The motif logos were then plotted by WebLogo (v3.7.12)71.

Revised version: Metagene analysis was performed using hg19 annotations according to previously reported tool, MetaplotR70. For motif analysis, the 20bp flanking sequence of each DRAM-seq editing site was extracted by Bedtools (v2.30.0)71. The motif logos were then plotted by WebLogo (v3.7.12)72.

Replicates analysis

Independent biological replicates of DRAM-ABE or DRAM-CBE in DRAM-seq analysis were separately compared by computing the Pearson correlation coefficient between the number of C-to-U mutations per mRNA between any two replicate experiments.

GO and KEGG analysis

Earlier version: GO and KEGG analysis of DRAM-seq edited mRNAs was performed using the DAVID bioinformatic database 72. GO terms with P value of less than 0.05 were considered as statistically significant.

Revised version: GO and KEGG analysis of DRAM-seq edited mRNAs was performed using the DAVID bioinformatic database 73. GO terms with a P value of less than 0.05 were considered statistically significant.

Added in the revised version: RNA pulldown assays

Added in the revised version: RNAs containing C or m5C residues and proteins expressing DRAM or DRAMmut were prepared in advance, processed using the Pierce™ Magnetic RNA-Protein Pull-Down Kit (Thermo) and finally analysed by Western blotting.

Statistical analysis

All data are expressed as mean ± S.E.M of three independent determinations. Data were analyzed through a two-tailed t-test. A probability of P < 0.05 was considered statistically significant; [, P < 0.05, [[, P < 0.01, *, P < 0.05, **, P <0.01, ***, P < 0.001 and ****, P < 0.0001 denote the significance thresholds; ns denotes not significant.